<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Lab:-Tower-of-Hanoi" data-toc-modified-id="Lab:-Tower-of-Hanoi-1">Lab: Tower of Hanoi</a></span></li><li><span><a href="#By-The-End-Of-This-Lab-You-Should-Be-Able-To" data-toc-modified-id="By-The-End-Of-This-Lab-You-Should-Be-Able-To-2">By The End Of This Lab You Should Be Able To</a></span></li><li><span><a href="#Code" data-toc-modified-id="Code-3">Code</a></span></li><li><span><a href="#Hints" data-toc-modified-id="Hints-4">Hints</a></span></li><li><span><a href="#Tests" data-toc-modified-id="Tests-5">Tests</a></span></li></ul></div>

Lab: Tower of Hanoi
-----

By The End Of This Lab You Should Be Able To
----

- Solve the Tower of Hanoi game with recursion
- Submit lab solutions with nbgrader

Code
-----

```python
reset -fs
```

<!-- grade_id: cell-d89c48480247256d, points: 0, grade: false, solution: true -->
```python
"Write a function to Tower of Hanoi with recursion"

def hanoi(n_disks=3, from_rod='left', to_rod='right', aux_rod='middle'): 
    "Solve Tower of Hanoi recursively"
    
    # Print state for autograding; DO NOT DELETE
    print(n_disks)
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return "Tower of Hanoi solved 🗼🙂"

```

Hints
-----

-----

Brian's solution is 3 lines (not including comments and print statements). 

Your solution can be longer.

----

[Python Tutor can help you visualize the state of your code](http://www.pythontutor.com/visualize.html#mode=edit)

-----

If you put this print statement in your function:

```python
print(f"Move disk {n_disks} from {from_rod} rod to {to_rod} rod.")
```

Your output should be something like this:

```
hanoi(n_disks=1)
Move disk 1 from left rod to right rod.

hanoi(n_disks=2)
Move disk 1 from left rod to middle rod.
Move disk 2 from left rod to right rod.
Move disk 1 from middle rod to right rod.

hanoi(n_disks=3)
Move disk 1 from left rod to right rod.
Move disk 2 from left rod to middle rod.
Move disk 1 from right rod to middle rod.
Move disk 3 from left rod to right rod.
Move disk 1 from middle rod to left rod.
Move disk 2 from middle rod to right rod.
Move disk 1 from left rod to right rod.
```

__Remember to remove this print statement before submitting assignment__

Tests
-----

<!-- grade_id: cell-be8a34383a4857ce, points: 3, grade: true, solution: false -->
```python
"""
Test code for hanoi function.
3 points for using recursion.
This cell should NOT give any errors when it is run.
"""

from bdb import Bdb
import sys

class RecursionDetected(Exception):
    pass

class RecursionDetector(Bdb):
    "From: https://stackoverflow.com/questions/36662181/is-there-a-way-to-check-if-function-is-recursive-in-python"
    def do_clear(self, arg):
        pass

    def __init__(self, *args):
        Bdb.__init__(self, *args)
        self.stack = set()

    def user_call(self, frame, argument_list):
        code = frame.f_code
        if code in self.stack:
            raise RecursionDetected
        self.stack.add(code)

    def user_return(self, frame, return_value):
        self.stack.remove(frame.f_code)

def test_recursion(func):
    detector = RecursionDetector()
    detector.set_trace()
    try:
        func()
    except RecursionDetected:
        return True
    else:
        return False
    finally:
        sys.settrace(None)
        
# Test if function is recursive
assert test_recursion(lambda: hanoi(n_disks=1))
assert test_recursion(lambda: hanoi(n_disks=2))
assert test_recursion(lambda: hanoi(n_disks=3))
```

<!-- grade_id: cell-811a7b634c7694da, points: 6, grade: true, solution: false -->
```python
"""
Test code for hanoi function.
6 points for optimal number of steps.
This cell should NOT give any errors when it is run.
"""

from io import StringIO
import sys

class Capturing(list):
    "From: https://stackoverflow.com/questions/16571150/how-to-capture-stdout-output-from-a-python-function-call"
    
    def __enter__(self):
        self._stdout = sys.stdout
        sys.stdout = self._stringio = StringIO()
        return self
    
    def __exit__(self, *args):
        self.extend(self._stringio.getvalue().splitlines())
        del self._stringio    # free up some memory
        sys.stdout = self._stdout


# n_disks=1        
with Capturing() as output:
    hanoi(n_disks=1)
    
assert output == ['1', '0', '0'] 
"""Note: 
There are extra moves with "0" discs.
We could prevent the zeros by adding a conditional that prevents their execution. That conditional is not required for this lab.
I want to focus just on the recursive logic. Those zeros are clues to the recursive logic required to solve this game.
"""
# Test for optimality after filtering out zeros
assert len([m for m in output if m != '0']) == 1

# n_disks=2
with Capturing() as output:
    hanoi(n_disks=2)
assert output == ['2', '1', '0', '0', '1', '0', '0']
# Test for optimality after filtering out zeros
assert len([m for m in output if m != '0']) == 3

# n_disks=3
with Capturing() as output:
    hanoi(n_disks=3)
assert output == ['3', '2', '1', '0', '0', '1', '0', '0', '2', '1', '0', '0', '1', '0', '0']
# Test for optimality after filtering out zeros
assert len([m for m in output if m != '0']) == 7
```

<br>
<br> 
<br>

----