<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#CartPole,-aka-Inverted-Pendulum" data-toc-modified-id="CartPole,-aka-Inverted-Pendulum-1">CartPole, aka Inverted Pendulum</a></span></li><li><span><a href="#OpenAI's-CartPole-Environment" data-toc-modified-id="OpenAI's-CartPole-Environment-2">OpenAI's CartPole Environment</a></span></li><li><span><a href="#Define-Environment" data-toc-modified-id="Define-Environment-3">Define Environment</a></span></li><li><span><a href="#Define-Neural-Network-" data-toc-modified-id="Define-Neural-Network--4">Define Neural Network </a></span></li><li><span><a href="#Define-Agent" data-toc-modified-id="Define-Agent-5">Define Agent</a></span></li><li><span><a href="#Train-the-model" data-toc-modified-id="Train-the-model-6">Train the model</a></span></li><li><span><a href="#Test-the-model" data-toc-modified-id="Test-the-model-7">Test the model</a></span></li><li><span><a href="#Grading-Submission-Notes" data-toc-modified-id="Grading-Submission-Notes-8">Grading Submission Notes</a></span></li><li><span><a href="#Bonus-Material" data-toc-modified-id="Bonus-Material-9">Bonus Material</a></span></li></ul></div>

<center><h2>CartPole, aka Inverted Pendulum</h2></center>
<br>
<center><img src="https://upload.wikimedia.org/wikipedia/commons/thumb/0/00/Cart-pendulum.svg/300px-Cart-pendulum.svg.png" width="35%"/></center>
<br><br><br>
<center><a href="https://fluxml.ai/experiments/cartPole/">Demo!</a></center>



<center><h2>OpenAI's CartPole Environment</h2></center>

A pole is attached to a cart by an un-actuated joint, and the cart moves along a frictionless track. 

The system is controlled by applying a force of +1 or -1 to the cart. 

The pendulum starts upright, and the goal is to prevent it from falling over. 

A reward of +1 is provided for every time-step that the pole remains upright. 

The episode ends when:

- The pole is more than 15 degrees from the vertical.
- The cart moves more than 2.4 units from the center.
- 200 time-steps.


```python
reset -fs
```

```python
import warnings
warnings.filterwarnings('ignore')
```

Define Environment
----

```python
# Import OpenAI's gym (easy way or hard way)
try:
    import gym 
except ImportError:
    import pip
    import sys
    import subprocess
    subprocess.check_call([sys.executable, '-m', 'pip', 'install', 'gym'])
    
    import gym
```

```python
env = gym.make('CartPole-v0')
```

```python
# Let's see how the RL problem is formulated
print('State size:  ', env.observation_space.shape[0] )
print('Action size: ', env.action_space.n)
```

Output:
```
State size:   4
Action size:  2
```

Let's read the docs:
https://github.com/openai/gym/wiki/CartPole-v0

Let's read the code: https://github.com/openai/gym/blob/master/gym/envs/classic_control/cartpole.py

Define Neural Network 
-----

```python
import numpy as np

from keras.models import Sequential
from keras.layers import Dense, Flatten
from keras.optimizers import Adam
```

```python
# Define a sample model
model = Sequential()
model.add(Flatten(input_shape=(1,) + env.observation_space.shape))
model.add(Dense(24, activation='relu'))
model.add(Dense(env.action_space.n, activation='linear'))
model.compile(loss='mse', optimizer=Adam(lr=0.01))
```

<!-- grade_id: cell-2231bf938a827285, points: 0, grade: true, solution: true -->
```python
# Define your model

# YOUR CODE HERE
# raise NotImplementedError()

model.summary()
```

Output:
```
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
flatten_6 (Flatten)          (None, 4)                 0         
_________________________________________________________________
dense_11 (Dense)             (None, 24)                120       
_________________________________________________________________
dense_12 (Dense)             (None, 2)                 50        
=================================================================
Total params: 170
Trainable params: 170
Non-trainable params: 0
_________________________________________________________________
```

Define Agent
----

```python
# Import keras-rl's gym (easy way or hard way)
try:
    import rl
except ImportError:
    import pip
    import sys
    import subprocess
    subprocess.check_call([sys.executable, '-m', 'pip', 'install', 'keras-rl'])
    import rl
```

```python
# Define memory
from rl.memory import SequentialMemory

memory = SequentialMemory(limit=50_000, 
                          window_length=1)
```

```python
# Define sample policy

from rl.policy import GreedyQPolicy

policy = GreedyQPolicy()
```

<!-- grade_id: cell-e3d15f2a794ccbad, points: 0, grade: true, solution: true -->
```python
# Define your policy

# YOUR CODE HERE
# raise NotImplementedError()

```

```python
# Define a sample agent
from rl.agents.dqn import DQNAgent

dqn = DQNAgent(model=model, 
               nb_actions=env.action_space.n, 
               memory=memory,
               nb_steps_warmup=15,
               target_model_update=1e-2,
               policy=policy)

dqn.compile(Adam(lr=1e-3), metrics=['mae'])
```

<!-- grade_id: cell-dabdbe3e55483280, points: 0, grade: true, solution: true -->
```python
# Define your agent

# YOUR CODE HERE
# raise NotImplementedError()

```

Train the model
----

```python
# Train model
dqn.fit(env, 
        nb_steps=1_500, 
        visualize=False, 
        verbose=2 # Episode logging
       )
```

Output:
```
Training for 1500 steps ...
   10/1500: episode: 1, duration: 0.501s, episode steps: 10, steps per second: 20, episode reward: 10.000, mean reward: 1.000 [1.000, 1.000], mean action: 1.000 [1.000, 1.000], mean observation: -0.155 [-3.091, 1.935], loss: --, mean_absolute_error: --, mean_q: --
   20/1500: episode: 2, duration: 2.538s, episode steps: 10, steps per second: 4, episode reward: 10.000, mean reward: 1.000 [1.000, 1.000], mean action: 1.000 [1.000, 1.000], mean observation: -0.142 [-3.029, 1.951], loss: 0.526494, mean_absolute_error: 0.530830, mean_q: 0.384968
   33/1500: episode: 3, duration: 0.072s, episode steps: 13, steps per second: 180, episode reward: 13.000, mean reward: 1.000 [1.000, 1.000], mean action: 0.846 [0.000, 1.000], mean observation: -0.139 [-2.931, 1.722], loss: 0.418333, mean_absolute_error: 0.468051, mean_q: 0.473482
   42/1500: episode: 4, duration: 0.060s, episode steps: 9, steps per second: 150, episode reward: 9.000, mean reward: 1.000 [1.000, 1.000], 
```

Test the model
----

```python
# Test model
test_results = dqn.test(env, nb_episodes=11, visualize=False)
```

Output:
```
Testing for 11 episodes ...
Episode 1: reward: 9.000, steps: 9
Episode 2: reward: 10.000, steps: 10
Episode 3: reward: 10.000, steps: 10
Episode 4: reward: 9.000, steps: 9
Episode 5: reward: 10.000, steps: 10
Episode 6: reward: 10.000, steps: 10
Episode 7: reward: 8.000, steps: 8
Episode 8: reward: 9.000, steps: 9
Episode 9: reward: 9.000, steps: 9
Episode 10: reward: 9.000, steps: 9
Episode 11: reward: 10.000, steps: 10
```

```python
# The max is 200 steps per eposide.
# The goal of the assignment to train an agent that performs about 180 (on average).

from statistics import mean

# Remove worst run
test_results.history['episode_reward'].remove(min(test_results.history['episode_reward']))

# Take the average the remaining runs
test_performance = mean(test_results.history['episode_reward'])

print(f"There current model gets {test_performance:.2f} steps.")
```

Output:
```
There current model gets 9.50 steps.
```

<!-- grade_id: cell-ba1d196810b21f40, points: 5, grade: true, solution: false -->
```python
# 5 points for over 100 steps per eposide.

assert test_performance > 100.00
```

Output:
```
ERROR:root:Internal Python error in the inspect module.
Below is the traceback from this internal error.

Traceback (most recent call last):
  File "/Users/brian/anaconda3/envs/rl-course/lib/python3.7/site-packages/IPython/core/interactiveshell.py", line 2878, in run_code
    exec(code_obj, self.user_global_ns, self.user_ns)
  File "<ipython-input-99-e353aa7f68bd>", line 3, in <module>
    assert test_performance > 100.00
AssertionError

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/Users/brian/anaconda3/envs/rl-course/lib/python3.7/site-packages/IPython/core/interactiveshell.py", line 1823, in showtraceback
    stb = value._render_traceback_()
AttributeError: 'AssertionError' object has no attribute '_render_traceback_'

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/Users/brian/anaconda3/envs/rl-course/lib/python3.7/site-packages/IPython/core/ultratb.py
```

<!-- grade_id: cell-ca8771487a5c5a1f, points: 5, grade: true, solution: false -->
```python
# 5 points for over 125 steps per eposide.

assert test_performance > 125.00
```

<!-- grade_id: cell-71810e3bdbc554a0, points: 10, grade: true, solution: false -->
```python
# 10 points for over 170 steps per eposide.

assert test_performance > 170.00
```

Grading Submission Notes
-------

If there is output, we'll grade your submitted lab without running the notebook. If there is __not__ output, we'll run the notebook to get output to grade.

It would behove you to submit a notebook with output.

Bonus Material
-----

Learn more about CartPole from the physics and control-model perspective  [here](https://danielpiedrahita.wordpress.com/portfolio/cart-pole-control/)