<center><h2>Lab: Rock, Paper, Scissors</h2></center>

<center><h2>Learning Outcomes</h2></center>

#### By the end of this session, you should be able to:

- Frame "Rock, Paper, Scissors" as a Reinforcement Learning problem.
- Implement several strategies for "Rock, Paper, Scissors".
- Discover a winning strategy to "Rock, Paper, Scissors" with multi-arm bandits.

Rock, Paper, Scissors as a Reinforcement Learning problem
--------

For Rock, Paper, Scissors, answer the following questions in the following cell:

1. Is it sequential?
1. What is the environment? 
1. Who is the agent?
1. What are the states?
1. What are the rewards?

5 points

<!-- grade_id: cell-717ff8136dd78454, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

Setup
-----

```python
reset -fs
```

```python
from collections import Counter
from random import choices, choice
from itertools import chain, cycle
from pprint import pprint

# Game Definition 
options = ['R', 'P', 'S'] # Rock, Paper, Scissor

scorer = dict(SP=1, PR=1, RS=1, 
              PS=-1, RP=-1, SR=-1, 
              SS=0, PP=0, RR=0)

# Deterministic decisions to win
ideal_response = {'P': 'S', # Scissors cuts Paper
                  'R': 'P', # Paper covers Rock
                  'S': 'R'} #  Rock crushes Scissors

# Strategy utilities 
def select_proportional(events, baseline=()):
    "Given a history of events, sample next event based on past frequnecy."
    rel_freq = Counter(chain(baseline, events))
    population, weights = zip(*rel_freq.items())
    return choices(population, weights)[0]

def select_maximum(events, baseline=()):
    "Given a history of events, select next event as the most common event seen so far."
    rel_freq = Counter(chain(baseline, events))
    return rel_freq.most_common(1)[0][0]

assert select_maximum('R R S'.split(), options) == 'R'
```

Strategies
----

<!-- grade_id: cell-d89c48480247256d, points: 0, grade: false, solution: true -->
```python
"""Write a function to randomly select among the possible options.

Hints:
- Read the rest of the lab before writing this function.
- Each function has the same parameters so there is a consistent API. Not all functions use all parameters.
- Feel free to use global functions and variables.
- It should be a simple one-liner.
"""

def random_reply(p1_history, p2_history):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-308b6a720f3853a4, points: 1, grade: true, solution: false -->
```python
"""
Test code for random_reply function.
1 point.
This cell should NOT give any errors when it is run.
"""

# Simulate plays
n = 1_000_000
outcomes = [random_reply(p1_history=None, p2_history=None) for _ in range(n)]

# Evaluate outcomes
c = Counter(outcomes)
assert set(c.keys()) == {'R', 'S', 'P'}
assert sum(c.values()) == n

# See if all outcomes are about equally likely
precision = -4
assert int(round(c['R'], precision)) == int(round(c['P'], precision))  == int(round(c['S'], precision)) 
```

<!-- grade_id: cell-97519a0b575fc09c, points: 0, grade: false, solution: true -->
```python
"""Write a function to select greedy.

Hint: Pick the right function with right parameters to generate prediction.
"""

def single_event_greedy(p1_history, p2_history):
    "When opponent plays R more than P or S, always respond with P."
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return ideal_response[prediction]
```

<!-- grade_id: cell-c6ec4a22b0e804d6, points: 2, grade: true, solution: false -->
```python
"""
Test code for single_event_greedy function.
2 points.
This cell should NOT give any errors when it is run.
"""

histories = [['R'],
             ['R', 'R'],
             ['R', 'R', 'P'],
             ['R', 'R', 'P', 'S'],
            ]
n = 100_000

for history in histories:
    outcomes = [single_event_greedy(p1_history=None, p2_history=history) for _ in range(n)]
    c = Counter(outcomes)

    assert (c['P'] == n) and (c['S'] == 0) and (c['R'] == 0)
```

<!-- grade_id: cell-3139ff18f6b4fca3, points: 0, grade: false, solution: true -->
```python
"""Write a function to select propotionally.

Hint: Pick the right function with right parameters to generate prediction.
"""

def single_event_proportional(p1_history, p2_history):
    "When opponent plays R two-thirds of the time, respond with P two-thirds of the time."
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return ideal_response[prediction]
```

<!-- grade_id: cell-cd5ad46a62538890, points: 2, grade: true, solution: false -->
```python
"""
Test code for single_event_proportional function.
2 points.
This cell should NOT give any errors when it is run.
"""

histories = [['R'],
             ['R', 'R'],
             ['R', 'R', 'R'],
            ]
n = 100_000

for history in histories:
    outcomes = [single_event_proportional(p1_history=None, p2_history=history) for _ in range(n)]
    c = Counter(outcomes)

    assert (c['P'] > 0) and (c['S'] > 0) and (c['R'] > 0)
    assert c['P']*2 > c['S']
    assert c['P']*2 > c['R']
```

<!-- grade_id: cell-876b92744b196086, points: 0, grade: false, solution: true -->
```python
"""Write a function to digraph_event_greedy.

Each line has been started for you.
"""

def digraph_event_greedy(p1_history, p2_history):
    """When opponent's most recent play is S and they usually play R two-thirds of the time after an S, 
    respond with P all of the time."""
    recent_play = None # TODO: Reassign this variable below to the correct answer.
    digraphs = None    # TODO: Reassign this variable below to the correct answer.
    followers = None   # TODO: Reassign this variable below to the correct answer.
    prediction = None  # TODO: Reassign this variable below to the correct answer.
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return ideal_response[prediction]
```

<!-- grade_id: cell-f57c09d8ec8e0197, points: 4, grade: true, solution: false -->
```python
"""
Test code for digraph_event_greedy function.
4 points.
This cell should NOT give any errors when it is run.
"""

histories = [['R'],
             ['R', 'R'],
             ['R', 'R', 'P'],
             ['R', 'R', 'P', 'S'],
            ]
n = 100_000

for history in histories:
    outcomes = [digraph_event_greedy(p1_history=None, p2_history=history) for _ in range(n)]
    c = Counter(outcomes)

    assert (c['P'] == n) and (c['S'] == 0) and (c['R'] == 0)
```

<!-- grade_id: cell-66ca59c5f03925fc, points: 0, grade: false, solution: true -->
```python
"""Write a function to digraph_event_proportional.

Similar to digraph_event_greedy and single_event_proportional
"""

def digraph_event_proportional(p1_history, p2_history):
    """When opponent's most recent play is S and they usually play R two-thirds of the time after an S, 
    respond with P two-thirds of the time.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return ideal_response[prediction]
```

<!-- grade_id: cell-4644b785ec195e6e, points: 4, grade: true, solution: false -->
```python
"""
Test code for digraph_event_proportional function.
4 point.
This cell should NOT give any errors when it is run.
"""

histories = [['R'],
             ['R', 'R'],
             ['R', 'R', 'R'],
            ]
n = 100_000

for history in histories:
    outcomes = [digraph_event_proportional(p1_history=None, p2_history=history) for _ in range(n)]
    c = Counter(outcomes)

    assert (c['P'] > 0) and (c['S'] > 0) and (c['R'] > 0)
    assert c['P']*2 > c['S']
    assert c['P']*2 > c['R']
```

Let's the Games Begin
----

```python
strategies = [random_reply, 
              single_event_greedy,
              single_event_proportional, 
              digraph_event_greedy,
              digraph_event_proportional]
```

```python
from typing import Callable, List

def play_and_learn(opposition: Callable,        # Opposition's strategy as a function
                   strategies: List=strategies, # Agent's possible strategies
                   trials: int=1_000, 
                   verbose: bool=False):
    "Simulate play and see how strategies empirically perform."
    strategy_range = range(len(strategies))
    weights = [1] * len(strategies)
    p1_history, p2_history = [], []
    cum_score = 0
    for trial in range(trials):
        # Choose our move based the relative weights
        our_moves = [strategy(p1_history, p2_history) for strategy in strategies]
        i = choices(strategy_range, weights)[0]
        our_move = our_moves[i]

        # Get opponent's move
        opponent_move = opposition(p2_history, p1_history)

        # Determine the winner
        score = scorer[our_move + opponent_move]
        if verbose:
            print(f'{our_move} ~ {opponent_move} = {score:+d}'
                  f'\t\t{strategies[i].__name__}')
        cum_score += score

        # Update move history and strategy weights
        p1_history.append(our_move)
        p2_history.append(opponent_move)
        for i, our_move in enumerate(our_moves):
            if scorer[our_move + opponent_move] == 1:
                weights[i] += 1
    print(f'Opposition strategy: {" ".join(opposition.__name__.title().split("_"))}')            
    print('Total score:', cum_score)
    print(f"\nRelative weights  | Our strategy")
    pprint(sorted([(weight, " ".join(strategy.__name__.title().split("_"))) for weight, strategy in zip(weights, strategies)]))
```

```python
def fixed_ratio(p1_history, p2_history):
    return choices(options, (1, 2, 3))[0]

play_and_learn(opposition=fixed_ratio)
```

```python
play_and_learn(opposition=random_reply)
```

```python
def cycling(series):
    iterator = cycle(series)
    def cycle_opponent(p1_history, p2_history):
        return next(iterator)
    return cycle_opponent

play_and_learn(opposition=cycling('RPRSS'))
```

```python
# If you want, you can play against computer. 
# Kinda boring - The computer is relentless.

# def human(p1_history, p2_history):
#     return input(f'Choose one of {options!r}: ')

# play_and_learn(opposition=human, trials=10, verbose=True)
```

Summarize what you learned about the strategies for Rock, Paper, Scissors in the next cell.

3 points.


<!-- grade_id: cell-7dce1e35b914d651, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

<br>
<br> 
<br>

----