<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Lab:-Gridworld" data-toc-modified-id="Lab:-Gridworld-1">Lab: Gridworld</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#By-the-end-of-this-session,-you-should-be-able-to:" data-toc-modified-id="By-the-end-of-this-session,-you-should-be-able-to:-2.0.1">By the end of this session, you should be able to:</a></span></li></ul></li></ul></li><li><span><a href="#(Grid)-World-Building" data-toc-modified-id="(Grid)-World-Building-3">(Grid) World Building</a></span></li><li><span><a href="#Markov-decision-process-(MDP)-Class" data-toc-modified-id="Markov-decision-process-(MDP)-Class-4">Markov decision process (MDP) Class</a></span></li><li><span><a href="#Value-Iteration-Function" data-toc-modified-id="Value-Iteration-Function-5">Value Iteration Function</a></span></li></ul></div>

Lab: Gridworld
-----

Learning Outcomes
----

#### By the end of this session, you should be able to:


- Model a gridworld in Python.
- Implement the Value Iteration function to calculate the values in a gridworld.

(Grid) World Building
-----

```python
reset -fs
```

<!-- grade_id: cell-b7cf369eee91207f, points: 0, grade: false, solution: true -->
```python
"""Create the create gridworld from lecture as a lists-of-lists.
Add reward values for each cell.
Use `None` for an obstacle / unreachable state. 
"""

warehouse = []

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bf5d4b9fdbe3856b, points: 4, grade: true, solution: false -->
```python
"""
Test code for warehouse data structure.
4 points.
This cell should NOT give any errors when it is run.
"""

assert len(warehouse) == 3 # Number of rows
assert [len(row) for row in warehouse] == [4, 4, 4] # Number of columns in each row
assert ["Found a None" for row in warehouse if None in row] == ["Found a None"] # Check for a single None
assert [sum(row) for row in warehouse if None not in row] == [0.88, -0.16] # For complete rows
assert [max(row) for row in warehouse if None not in row] == [1, -0.04] # For complete rows
assert [sum(row) for row in list(map(list, zip(*warehouse))) if None not in row] == [-0.12, -0.12, -0.04] # For complete columns
assert [min(row) for row in list(map(list, zip(*warehouse))) if None not in row] == [-0.04, -0.04, -1] # For complete columns
```

Markov decision process (MDP) Class
------

A Markov decision process (MDP) class for a two-dimensional grid has been implemented for you.

Please read carefully through it

```python
class MarkovDecisionProcess():
    "A two-dimensional grid Markov decision process (MDP)."

    def __init__(self, grid, terminals, gamma=1):
        grid.reverse() # Row 0 is on bottom, not on top; The difference between gridworld and Python
        self.grid = grid
        self.terminals = terminals
        self.gamma = gamma
        self.n_rows = len(grid)
        self.n_cols = len(grid[0])

        states = set()
        reward = {}
        for x in range(self.n_cols):
            for y in range(self.n_rows):
                if grid[y][x]:
                    states.add((x, y))
                    reward[(x, y)] = grid[y][x]
                else:
                    reward[(x, y)] = None
        self.reward = reward
        self.states = states
        
        self.actlist = [(1, 0), (0, 1), (-1, 0), (0, -1)] # An action is an (x, y) unit vector; e.g. (1, 0) means move east.
        
        transitions = {}
        for s in states:
            transitions[s] = {}
            for a in self.actlist:
                transitions[s][a] = self.calculate_t(s, a)
        self.transitions=transitions
        
    def actions(self, state):
        """Return a list of actions that can be performed in this state. By default, a
        fixed list of actions, except for terminal states. Override this
        method if you need to specialize by state."""

        if state in self.terminals:
            return [None]
        else:
            return self.actlist

    def calculate_t(self, state, action):
        "Caculate the transition, including probability of not completing intended action."
        if action:
            return [(0.8, self.go(state, action)), # Intended
                    (0.1, self.go(state, self.turn_right(action))), # Error
                    (0.1, self.go(state, self.turn_left(action)))]  # Error
        else:
            return [(0.0, state)]
    
    def t(self, state, action):
        return self.transitions[state][action] if action else [(0.0, state)]
    
    def turn_heading(self, heading, inc, headings=[(1, 0), (0, 1), (-1, 0), (0, -1)]):
        return headings[(headings.index(heading) + inc) % len(headings)]

    def turn_right(self, heading):
        return self.turn_heading(heading, -1)

    def turn_left(self, heading):
        return self.turn_heading(heading, +1)

    def get_reward(self, state):
        "Return a numeric reward for this state."
        return self.reward[state]
 
    def go(self, state, direction):
        "Return the state that results from going in this direction."
        state1 = tuple(map(lambda x, y: x+y, state, direction))
        return state1 if state1 in self.states else state

```

<!-- grade_id: cell-9b30ea71c3de966f, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Value Iteration Function
----

<!-- grade_id: cell-c0e99e2c9f78f178, points: 0, grade: false, solution: true -->
```python
"""Write value iteration function to solve mdp.

See lecture notes for pseudocode.
"""

def value_iteration(mdp, epsilon=0.001)->dict:
    "Solve an MDP by value iteration."
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0257007adfa31dbb, points: 1, grade: true, solution: false -->
```python
"""
Test code for value_iteration function.
1 point for correctly formatting answer.
This cell should NOT give any errors when it is run.
"""

mdp_values = value_iteration(warehouse_mdp)
assert type(mdp_values) == dict
assert len(mdp_values.keys()) == 11
assert sorted(mdp_values.keys()) == [\
 (0, 0),
 (0, 1),
 (0, 2),
 (1, 0),
 (1, 2),
 (2, 0),
 (2, 1),
 (2, 2),
 (3, 0),
 (3, 1),
 (3, 2)
]
```

<!-- grade_id: cell-ac4f0b7d4152658b, points: 1, grade: true, solution: false -->
```python
"""
Test code for value_iteration function.
1 point for not changing terminal states.
This cell should NOT give any errors when it is run.
"""
mdp_values = value_iteration(warehouse_mdp)
assert mdp_values[(3, 1)] == -1.0
assert mdp_values[(3, 2)] == 1.0
```

<!-- grade_id: cell-569c829cb5cd6555, points: 1, grade: true, solution: false -->
```python
"""
Test code for value_iteration function.
1 point for NOT calcuting value for unreachable state.
This cell should NOT give any errors when it is run.
"""
mdp_values = value_iteration(warehouse_mdp)
try: 
    mdp_values[(1, 1)] # Look for key that should not be there
    assert False
except KeyError:
    assert True
```

<!-- grade_id: cell-3ec0b8efc5f2204c, points: 5, grade: true, solution: false -->
```python
"""
Test code for value_iteration function.
5 points for correct answer with generous rounding.
This cell should NOT give any errors when it is run.
"""

expected_values = {\
 (0, 0): 0.7053082191780823,
 (0, 1): 0.7615582191780823,
 (0, 2): 0.8115582191780822,
 (1, 0): 0.6553082191780822,
 (1, 2): 0.8678082191780823,
 (2, 0): 0.6114155251141552,
 (2, 1): 0.6602739726027398,
 (2, 2): 0.9178082191780822,
 (3, 0): 0.3879249112125825,
 (3, 1): -1.0,
 (3, 2): 1.0}

computed_values = dict(sorted(value_iteration(warehouse_mdp, epsilon=1).items())) # Sort solution for quick visual inspection
rounded_computed_values = {k: round(v, 2) for k,v in computed_values.items()}
rounded_expected_values = {k: round(v, 2) for k,v in expected_values.items()}
assert rounded_computed_values == rounded_expected_values
```

<!-- grade_id: cell-1f07d2a13bc3428d, points: 10, grade: true, solution: false -->
```python
"""
Test code for value_iteration function.
10 points for precisely correct answer.
This cell should NOT give any errors when it is run.
"""

expected_values = {\
 (0, 0): 0.7053082191780823,
 (0, 1): 0.7615582191780823,
 (0, 2): 0.8115582191780822,
 (1, 0): 0.6553082191780822,
 (1, 2): 0.8678082191780823,
 (2, 0): 0.6114155251141552,
 (2, 1): 0.6602739726027398,
 (2, 2): 0.9178082191780822,
 (3, 0): 0.3879249112125825,
 (3, 1): -1.0,
 (3, 2): 1.0}

computed_values = dict(sorted(value_iteration(warehouse_mdp, epsilon=1).items())) # Sort solution for quick visual inspection
assert rounded_computed_values == rounded_expected_values
```

<br>
<br> 
<br>

----